(assume h1 (! (not (! (or true false) :named @p_1)) :named @p_3))
(step t2 (cl (= @p_1 (! (or true) :named @p_2))) :rule or_simplify)
(step t3 (cl (= @p_2 true)) :rule or_simplify)
(step t4 (cl (= @p_1 true)) :rule trans :premises (t2 t3))
(step t5 (cl (= @p_3 (! (not true) :named @p_4))) :rule cong :premises (t4))
(step t6 (cl (= @p_4 false)) :rule not_simplify)
(step t7 (cl (! (= @p_3 false) :named @p_5)) :rule trans :premises (t5 t6))
(step t8 (cl (! (not @p_5) :named @p_7) (! (not @p_3) :named @p_6) false) :rule equiv_pos2)
(step t9 (cl (not @p_6) @p_1) :rule not_not)
(step t10 (cl @p_7 @p_1 false) :rule th_resolution :premises (t9 t8))
(step t11 (cl false) :rule th_resolution :premises (h1 t7 t10))
(step t12 (cl (not false)) :rule false)
(step t13 (cl) :rule resolution :premises (t11 t12))
