(assume h1 (not (or true false)))
(step t2 (cl (= (or true false) (or true))) :rule or_simplify)
(step t3 (cl (= (or true) true)) :rule or_simplify)
(step t4 (cl (= (or true false) true)) :rule trans :premises (t2 t3))
(step t5 (cl (= (not (or true false)) (not true))) :rule cong :premises (t4))
(step t6 (cl (= (not true) false)) :rule not_simplify)
(step t7 (cl (= (not (or true false)) false)) :rule trans :premises (t5 t6))
(step t8 (cl (not (= (not (or true false)) false)) (not (not (or true false))) false) :rule equiv_pos2)
(step t9 (cl (not (not (not (or true false)))) (or true false)) :rule not_not)
(step t10 (cl (not (= (not (or true false)) false)) (or true false) false) :rule th_resolution :premises (t9 t8))
(step t11 (cl false) :rule th_resolution :premises (h1 t7 t10))
(step t12 (cl (not false)) :rule false)
(step t13 (cl) :rule resolution :premises (t11 t12))
