(assume h1 (! (and true (! (not true) :named @p_1)) :named @p_2))
(step t2 (cl (= @p_1 false)) :rule not_simplify)
(step t3 (cl (= @p_2 (! (and true false) :named @p_3))) :rule cong :premises (t2))
(step t4 (cl (= @p_3 (! (and false) :named @p_4))) :rule and_simplify)
(step t5 (cl (= @p_4 false)) :rule and_simplify)
(step t6 (cl (! (= @p_2 false) :named @p_5)) :rule trans :premises (t3 t4 t5))
(step t7 (cl (not @p_5) (not @p_2) false) :rule equiv_pos2)
(step t8 (cl false) :rule th_resolution :premises (h1 t6 t7))
(step t9 (cl (not false)) :rule false)
(step t10 (cl) :rule resolution :premises (t8 t9))
