(assume h1 (not (=> (= y (+ x 1)) (= (f x) (f (- y 1))))))
(step t2 (cl (= (not (=> (= y (+ x 1)) (= (f x) (f (- y 1))))) (and (= y (+ x 1)) (not (= (f x) (f (- y 1))))))) :rule bool_simplify)
(step t3 (cl (not (= (not (=> (= y (+ x 1)) (= (f x) (f (- y 1))))) (and (= y (+ x 1)) (not (= (f x) (f (- y 1))))))) (not (not (=> (= y (+ x 1)) (= (f x) (f (- y 1)))))) (and (= y (+ x 1)) (not (= (f x) (f (- y 1)))))) :rule equiv_pos2)
(step t4 (cl (not (not (not (=> (= y (+ x 1)) (= (f x) (f (- y 1))))))) (=> (= y (+ x 1)) (= (f x) (f (- y 1))))) :rule not_not)
(step t5 (cl (not (= (not (=> (= y (+ x 1)) (= (f x) (f (- y 1))))) (and (= y (+ x 1)) (not (= (f x) (f (- y 1))))))) (=> (= y (+ x 1)) (= (f x) (f (- y 1)))) (and (= y (+ x 1)) (not (= (f x) (f (- y 1)))))) :rule th_resolution :premises (t4 t3))
(step t6 (cl (and (= y (+ x 1)) (not (= (f x) (f (- y 1)))))) :rule th_resolution :premises (h1 t2 t5))
(step t7 (cl (= (+ x 1) (+ 1 x))) :rule sum_simplify)
(step t8 (cl (= (= y (+ x 1)) (= y (+ 1 x)))) :rule cong :premises (t7))
(step t9 (cl (= (and (= y (+ x 1)) (not (= (f x) (f (- y 1))))) (and (= y (+ 1 x)) (not (= (f x) (f (- y 1))))))) :rule cong :premises (t8))
(step t10 (cl (not (= (and (= y (+ x 1)) (not (= (f x) (f (- y 1))))) (and (= y (+ 1 x)) (not (= (f x) (f (- y 1))))))) (not (and (= y (+ x 1)) (not (= (f x) (f (- y 1)))))) (and (= y (+ 1 x)) (not (= (f x) (f (- y 1)))))) :rule equiv_pos2)
(step t11 (cl (and (= y (+ 1 x)) (not (= (f x) (f (- y 1)))))) :rule th_resolution :premises (t6 t9 t10))
(step t12 (cl (= y (+ 1 x))) :rule and :premises (t11))
(step t13 (cl (not (= (f x) (f (- y 1))))) :rule and :premises (t11))
(step t14 (cl (not (= x (- y 1))) (= (f x) (f (- y 1)))) :rule eq_congruent)
(step t15 (cl (not (= x (- y 1)))) :rule resolution :premises (t14 t13))
(step t16 (cl (or (= x (- y 1)) (not (<= x (- y 1))) (not (<= (- y 1) x)))) :rule la_disequality)
(step t17 (cl (= x (- y 1)) (not (<= x (- y 1))) (not (<= (- y 1) x))) :rule or :premises (t16))
(step t18 (cl (not (<= x (- y 1))) (not (<= (- y 1) x))) :rule resolution :premises (t17 t15))
(step t19 (cl (<= (- y 1) x) (not (= y (+ 1 x)))) :rule la_generic)
(step t20 (cl (<= (- y 1) x)) :rule resolution :premises (t19 t12))
(step t21 (cl (not (<= x (- y 1)))) :rule resolution :premises (t18 t20))
(step t22 (cl (<= x (- y 1)) (not (= y (+ 1 x)))) :rule la_generic)
(step t23 (cl) :rule resolution :premises (t22 t12 t21))
