(assume a0 (not (or (not (= (op_3 op_2) op_0)) (not (op_1 (op_3 op_2))) (op_1 op_0))))
(step t1 (cl (not (= (not (or (not (= (op_3 op_2) op_0)) (not (op_1 (op_3 op_2))) (op_1 op_0))) (not (or (not (op_1 (op_3 op_2))) (op_1 (op_3 op_2)))))) (not (not (or (not (= (op_3 op_2) op_0)) (not (op_1 (op_3 op_2))) (op_1 op_0)))) (not (or (not (op_1 (op_3 op_2))) (op_1 (op_3 op_2))))) :rule equiv_pos2)
(step t2 (cl (= (= (op_3 op_2) op_0) (= op_0 (op_3 op_2)))) :rule hole :args (ARITH_POLY_NORM (= (= (op_3 op_2) op_0) (= op_0 (op_3 op_2)))))
(step t3 (cl (= (not (= (op_3 op_2) op_0)) (not (= op_0 (op_3 op_2))))) :rule cong :premises (t2))
(step t4 (cl (= (not (op_1 (op_3 op_2))) (not (op_1 (op_3 op_2))))) :rule refl)
(step t5 (cl (= (op_1 op_0) (op_1 op_0))) :rule refl)
(step t6 (cl (= (or (not (= (op_3 op_2) op_0)) (not (op_1 (op_3 op_2))) (op_1 op_0)) (or (not (= op_0 (op_3 op_2))) (not (op_1 (op_3 op_2))) (op_1 op_0)))) :rule cong :premises (t3 t4 t5))
(step t7 (cl (= (not (or (not (= (op_3 op_2) op_0)) (not (op_1 (op_3 op_2))) (op_1 op_0))) (not (or (not (= op_0 (op_3 op_2))) (not (op_1 (op_3 op_2))) (op_1 op_0))))) :rule cong :premises (t6))
(step t8 (cl (not (not (not (= op_0 (op_3 op_2))))) (= op_0 (op_3 op_2))) :rule not_not)
(step t9 (cl (not (= (not (or (not (= (op_3 op_2) op_0)) (not (op_1 (op_3 op_2))) (op_1 op_0))) (not (or (not (= op_0 (op_3 op_2))) (not (op_1 (op_3 op_2))) (op_1 op_0))))) (not (not (or (not (= (op_3 op_2) op_0)) (not (op_1 (op_3 op_2))) (op_1 op_0)))) (not (or (not (= op_0 (op_3 op_2))) (not (op_1 (op_3 op_2))) (op_1 op_0)))) :rule equiv_pos2)
(step t10 (cl (not (or (not (= op_0 (op_3 op_2))) (not (op_1 (op_3 op_2))) (op_1 op_0)))) :rule resolution :premises (t9 t7 a0))
(step t11 (cl (not (not (= op_0 (op_3 op_2))))) :rule not_or :premises (t10))
(step t12 (cl (= op_0 (op_3 op_2))) :rule resolution :premises (t8 t11))
(step t13 (cl (= (op_3 op_2) (op_3 op_2))) :rule refl)
(step t14 (cl (= (= op_0 (op_3 op_2)) (= (op_3 op_2) (op_3 op_2)))) :rule cong :premises (t12 t13))
(step t15 (cl (= (not (= op_0 (op_3 op_2))) (not (= (op_3 op_2) (op_3 op_2))))) :rule cong :premises (t14))
(step t16 (cl (= (not (op_1 (op_3 op_2))) (not (op_1 (op_3 op_2))))) :rule refl)
(step t17 (cl (= (op_1 op_0) (op_1 (op_3 op_2)))) :rule cong :premises (t12))
(step t18 (cl (= (or (not (= op_0 (op_3 op_2))) (not (op_1 (op_3 op_2))) (op_1 op_0)) (or (not (= (op_3 op_2) (op_3 op_2))) (not (op_1 (op_3 op_2))) (op_1 (op_3 op_2))))) :rule cong :premises (t15 t16 t17))
(step t19 (cl (= (not (or (not (= op_0 (op_3 op_2))) (not (op_1 (op_3 op_2))) (op_1 op_0))) (not (or (not (= (op_3 op_2) (op_3 op_2))) (not (op_1 (op_3 op_2))) (op_1 (op_3 op_2)))))) :rule cong :premises (t18))
(step t20 (cl (= (= (op_3 op_2) (op_3 op_2)) true)) :rule all_simplify :args (eq_simplify1 (op_3 op_2)))
(step t21 (cl (= (not (= (op_3 op_2) (op_3 op_2))) (not true))) :rule cong :premises (t20))
(step t22 (cl (= (not true) false)) :rule not_simplify)
(step t23 (cl (= (not (= (op_3 op_2) (op_3 op_2))) false)) :rule trans :premises (t21 t22))
(step t24 (cl (= (op_1 (op_3 op_2)) (op_1 (op_3 op_2)))) :rule refl)
(step t25 (cl (= (or (not (= (op_3 op_2) (op_3 op_2))) (not (op_1 (op_3 op_2))) (op_1 (op_3 op_2))) (or false (not (op_1 (op_3 op_2))) (op_1 (op_3 op_2))))) :rule cong :premises (t23 t4 t24))
(step t26 (cl (= (or false (not (op_1 (op_3 op_2))) (op_1 (op_3 op_2))) (or (not (op_1 (op_3 op_2))) (op_1 (op_3 op_2))))) :rule all_simplify :args (or_simplify2 (cvc5_nary_op) (cvc5_nary_op (not (op_1 (op_3 op_2))))))
(step t27 (cl (= (or (not (= (op_3 op_2) (op_3 op_2))) (not (op_1 (op_3 op_2))) (op_1 (op_3 op_2))) (or (not (op_1 (op_3 op_2))) (op_1 (op_3 op_2))))) :rule trans :premises (t25 t26))
(step t28 (cl (= (not (or (not (= (op_3 op_2) (op_3 op_2))) (not (op_1 (op_3 op_2))) (op_1 (op_3 op_2)))) (not (or (not (op_1 (op_3 op_2))) (op_1 (op_3 op_2)))))) :rule cong :premises (t27))
(step t29 (cl (= (not (or (not (= op_0 (op_3 op_2))) (not (op_1 (op_3 op_2))) (op_1 op_0))) (not (or (not (op_1 (op_3 op_2))) (op_1 (op_3 op_2)))))) :rule trans :premises (t19 t28))
(step t30 (cl (= (not (or (not (= (op_3 op_2) op_0)) (not (op_1 (op_3 op_2))) (op_1 op_0))) (not (or (not (op_1 (op_3 op_2))) (op_1 (op_3 op_2)))))) :rule trans :premises (t7 t29))
(step t31 (cl (not (or (not (op_1 (op_3 op_2))) (op_1 (op_3 op_2))))) :rule resolution :premises (t1 t30 a0))
(step t32 (cl (not (op_1 (op_3 op_2)))) :rule not_or :premises (t31))
(step t33 (cl (not (not (not (op_1 (op_3 op_2))))) (op_1 (op_3 op_2))) :rule not_not)
(step t34 (cl (not (not (op_1 (op_3 op_2))))) :rule not_or :premises (t31))
(step t35 (cl (op_1 (op_3 op_2))) :rule resolution :premises (t33 t34))
(step t36 (cl) :rule resolution :premises (t32 t35))
