(assume h1 (not (or (not (= (op_3 op_2) op_0)) (not (op_1 (op_3 op_2))) (op_1 op_0))))
(step t2 (cl (not (not (= (op_3 op_2) op_0)))) :rule not_or :premises (h1))
(step t3 (cl (not (not (not (= (op_3 op_2) op_0)))) (= (op_3 op_2) op_0)) :rule not_not)
(step t4 (cl (= (op_3 op_2) op_0)) :rule th_resolution :premises (t3 t2))
(step t5 (cl (not (not (op_1 (op_3 op_2))))) :rule not_or :premises (h1))
(step t6 (cl (not (not (not (op_1 (op_3 op_2))))) (op_1 (op_3 op_2))) :rule not_not)
(step t7 (cl (op_1 (op_3 op_2))) :rule th_resolution :premises (t6 t5))
(step t8 (cl (not (op_1 op_0))) :rule not_or :premises (h1))
(step t9 (cl (not (= (op_3 op_2) op_0)) (not (op_1 (op_3 op_2))) (op_1 op_0)) :rule eq_congruent_pred)
(step t10 (cl) :rule resolution :premises (t9 t4 t7 t8))
