(assume h1 (! (not (! (or op_0 (not op_0)) :named @p_1)) :named @p_2))
(step t2 (cl (= @p_1 true)) :rule or_simplify)
(step t3 (cl (= @p_2 (! (not true) :named @p_3))) :rule cong :premises (t2))
(step t4 (cl (= @p_3 false)) :rule not_simplify)
(step t5 (cl (! (= @p_2 false) :named @p_4)) :rule trans :premises (t3 t4))
(step t6 (cl (! (not @p_4) :named @p_6) (! (not @p_2) :named @p_5) false) :rule equiv_pos2)
(step t7 (cl (not @p_5) @p_1) :rule not_not)
(step t8 (cl @p_6 @p_1 false) :rule th_resolution :premises (t7 t6))
(step t9 (cl false) :rule th_resolution :premises (h1 t5 t8))
(step t10 (cl (not false)) :rule false)
(step t11 (cl) :rule resolution :premises (t9 t10))
