(assume a0 (not (or true false)))
(step t1 (cl (not (= (not (or true false)) false)) (not (not (or true false))) false) :rule equiv_pos2)
(step t2 (cl (= (or true false) true)) :rule all_simplify :args (evaluate))
(step t3 (cl (= (not (or true false)) (not true))) :rule cong :premises (t2))
(step t4 (cl (= (not true) false)) :rule all_simplify :args (evaluate))
(step t5 (cl (= (not (or true false)) false)) :rule trans :premises (t3 t4))
(step t6 (cl) :rule resolution :premises (t1 t5 a0))
