(assume a0 (and true (not true)))
(step t1 (cl (not (= (and true (not true)) false)) (not (and true (not true))) false) :rule equiv_pos2)
(step t2 (cl (= (and true (not true)) false)) :rule all_simplify :args (evaluate))
(step t3 (cl (= (not true) false)) :rule all_simplify :args (evaluate))
(step t4 (cl (= false (not true))) :rule symm :premises (t3))
(step t5 (cl (= (and true (not true)) (not true))) :rule trans :premises (t2 t4))
(step t6 (cl (= (not true) false)) :rule all_simplify :args (evaluate))
(step t7 (cl (= (and true (not true)) false)) :rule trans :premises (t5 t6))
(step t8 (cl) :rule resolution :premises (t1 t7 a0))
