(assume h1 (and (or op_0 op_2 op_1) (or (not op_0) (not op_2) (not op_1)) (or (not op_0) op_2) (or (not op_2) op_1) (or (not op_1) op_0)))
(step t2 (cl (or op_0 op_2 op_1)) :rule and :premises (h1))
(step t3 (cl op_0 op_2 op_1) :rule or :premises (t2))
(step t4 (cl (or (not op_0) (not op_2) (not op_1))) :rule and :premises (h1))
(step t5 (cl (not op_0) (not op_2) (not op_1)) :rule or :premises (t4))
(step t6 (cl (or (not op_0) op_2)) :rule and :premises (h1))
(step t7 (cl (not op_0) op_2) :rule or :premises (t6))
(step t8 (cl (or (not op_2) op_1)) :rule and :premises (h1))
(step t9 (cl (not op_2) op_1) :rule or :premises (t8))
(step t10 (cl (or (not op_1) op_0)) :rule and :premises (h1))
(step t11 (cl (not op_1) op_0) :rule or :premises (t10))
(step t12 (cl op_0) :rule resolution :premises (t3 t9 t11))
(step t13 (cl op_2) :rule resolution :premises (t7 t12))
(step t14 (cl (not op_1)) :rule resolution :premises (t5 t13 t12))
(step t15 (cl) :rule resolution :premises (t9 t13 t14))
