(assume h1 (not (or p (not p))))
(step t2 (cl (= (or p (not p)) true)) :rule or_simplify)
(step t3 (cl (= (not (or p (not p))) (not true))) :rule cong :premises (t2))
(step t4 (cl (= (not true) false)) :rule not_simplify)
(step t5 (cl (= (not (or p (not p))) false)) :rule trans :premises (t3 t4))
(step t6 (cl (not (= (not (or p (not p))) false)) (not (not (or p (not p)))) false) :rule equiv_pos2)
(step t7 (cl (not (not (not (or p (not p))))) (or p (not p))) :rule not_not)
(step t8 (cl (not (= (not (or p (not p))) false)) (or p (not p)) false) :rule th_resolution :premises (t7 t6))
(step t9 (cl false) :rule th_resolution :premises (h1 t5 t8))
(step t10 (cl (not false)) :rule false)
(step t11 (cl) :rule resolution :premises (t9 t10))
